Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ilya Grebenkin"
COLLABORATORS = "-"
```

---

# Large-scale eigenvalue problem with structured matrices


Consider a diagonal matrix $D$ of the size $n\times n$ and a column vector $u$ of size $n$. Consider a matrix 

$$
A = D +  \frac{u\, u^T}{u^T u}
$$

Drawing the elements of $D$ and $u$ from a standard normal distribution, find the smallest (by the absolute value) eigenvalue of $A$. You will need to consider values of up to $n=10^5$. Note that matrices of this size do not fit into memory, so you will need to write a matrix-free implementation which uses the special structure of $A$. You may want to consider using the Sherman-Morrison formula and the so-called [Bunch-Nielsen-Sorensen formula](https://en.wikipedia.org/wiki/Bunch%E2%80%93Nielsen%E2%80%93Sorensen_formula).

```python
import numpy as np
from numpy.testing import assert_allclose
from numpy.linalg import eigvals, eig
```

```python
from math import sqrt

def generate_d_u(n, rndm):
    """Generate two vectors of size $n$ drawn from a standard normal distribution."""
    d = rndm.standard_normal(size=n)
    u = rndm.standard_normal(size=n)
    u /= sqrt(u @ u)
    return d, u
```

We start from a simple case of small $n$, and write out the matrices explicitly.

```python
rndm = np.random.default_rng(12345)
d, u = generate_d_u(5, rndm)

A = np.diag(d) + u[:, None] @ u[None, :]
A
```

Output:
```
array([[-1.34292291,  0.14935839, -0.07085692, -0.03942634,  0.21324611],
       [ 0.14935839,  1.53946816, -0.13081332, -0.07278739,  0.39368675],
       [-0.07085692, -0.13081332, -0.80860276,  0.03453097, -0.18676843],
       [-0.03942634, -0.07278739,  0.03453097, -0.23995945, -0.10392203],
       [ 0.21324611,  0.39368675, -0.18676843, -0.10392203,  0.4867421 ]])
```

## Find the min eigenvalue of $A$

Min eigenvalue of $A$ is the max eigenvalue of $A^{-1}$. Use the power iteration with $A^{-1}$ as the iteration matrix.

Demo version: use full matrices for clarity. Also we only do a fixed number of iterations instead of a proper convergence check.

```python
def max_eigval(a, vec, eps=1e-3, maxiter=1000):
    """Реализует алгоритм степенной итерации поиска максимального собственного значения матрицы `a` 
        с начальным приближением собственного вектора vec.
    
    Parameters
    ----------
    a : array-like of floats, shape (n, n)
        Введённая матрица
    vec : array-like of floats, shape (n,)
        Введённое начальное приближение собственного вектора
    eps : float
        Введённая абсолютная точность
    maxiter : integer
        Максимальное количество итераций
    
    Returns
    -------
    lm : float
        Приближение максимального собственного значения
    outvec : array of floats, shape (n,)
        Приближение собственного вектора
    iters : integer
        Количество итераций
    """
    a = np.asarray(a, dtype=float)
    if a.ndim != 2:
        raise ValueError("a.ndim = %s, expected 2" % a.ndim)
    
    vec = np.asarray(vec, dtype=float)
    if vec.ndim != 1:
        raise ValueError("vec.ndim = %s, expected 1" % vec.ndim)
        
    if (a.shape[0] != a.shape[1]):
        raise ValueError("a is not square")
        
    if (a.shape[0] != vec.shape[0]):
        raise ValueError("vec should be of length %s" % a.shape[0])
    
    # YOUR CODE HERE
    x_k = np.copy(vec)
    iters = 0
    while iters < maxiter:
        iters += 1
        y_k1 = a @ x_k.T
        x_k1 = y_k1 / np.linalg.norm(y_k1)
        lambda_k = x_k @ y_k1.T
        if abs(np.linalg.norm(a @ x_k1) - np.linalg.norm(lambda_k * x_k1)) <= eps:
            break
        x_k = np.copy(x_k1)
    return lambda_k, x_k1, iters
```

<!-- grade_id: cell-d833d5a69f14bd61, points: 0, grade: false, solution: true -->
```python
def power_iter(A, x0, niter=20):
    """Compute min eigenvalue and corresponding eigenvector of matrix A.
    
    Use the power iteration.
    
    Parameters
    ----------
    A : ndarray of floats, shape (n, n)
        original matrix
    x0 : ndarray of floats, shape (n,)
        The starting vector for the iteration
    niter : int, optional
        the number of iterations (default is 20)
    
    Returns
    -------
    la : float
        min eigenvalue
    x : ndarray of floats, shape (n,)
        the iteration vector on the last iteration
    """
    # YOUR CODE HERE

    A_inv = np.linalg.inv(A)
    
    la, x0, iters = max_eigval(A_inv, x0, 1e-6, maxiter=niter)
    
    return 1/la, x0
```

```python
d, u = generate_d_u(5, rndm=np.random.default_rng(1234))
A = np.diag(d) + u[:, None] @ u[None, :]

la, x = power_iter(A, x0=np.ones(5))

# Check that the iteration vector approximates the eigenvector
np.linalg.norm(la * x - A @ x)
```

Output:
```
9.900473222338938e-06
```

<!-- grade_id: cell-8d3c475509476796, points: 3, grade: true, solution: false -->
```python
# this is a graded test

d1, u1 = generate_d_u(6, rndm)
A1 = np.diag(d1) + u1[:, None] @ u1[None, :]

la1, x1 = power_iter(A1, x0=np.ones(6))

err1 = np.linalg.norm(la1 * x1 - A1 @ x1)
assert_allclose(err1, 0., atol=1e-4)

```

## Check the BNS formula

Implement the [Bunch-Nielsen-Sorensen formula](https://en.wikipedia.org/wiki/Bunch%E2%80%93Nielsen%E2%80%93Sorensen_formula) for the matrix of the above structure. 

<!-- grade_id: cell-4ea65883045a7664, points: 0, grade: false, solution: true -->
```python
def BNS(d, u, la):
    """Compute the eigenvector of $D + u u^T$ for a given eigenvalue.
    
    Given the matrix $A = D + u u^T$ where $D$ is a diagonal matrix and
    $u$ is a 1D vector, and also a precomputed eigenvalue of $A$,
    compute its corresponding eigenvector using the B-N-S formula.
    
    Parameters
    ----------
    d : ndarray, shape (n,)
        Diagonal elements of the matrix $D$
    u : ndarray, shape (n,)
        1D array representing the vector $u$
    la : float
        An eigenvalue of $A = D + u u^T$
        
    Returns
    -------
    q : ndarray, shape (n,)
        The normalized eigenvector of $A$, which corresponds
        to the eigenvalue `la`.
    """
    # YOUR CODE HERE
    

    I = np.eye(x0.shape[0])
    
#     lambda_k, x0, iters = max_eigval(A, x0, 1e-3, niter)
#     sigma = lambda_k + 0.1
#     A = A - sigma * I
#     lambda_k, x0, iters = max_eigval(A, x0, 1e-3, niter)
#     la = lambda_k - sigma
    
    x_k = np.copy(vec)
    iters = 0
    while np.linalg.norm(x_k1 - x_k) < 1e-4:
        iters += 1
#         y_k1 = a @ x_k.T
#         x_k1 = y_k1 / np.linalg.norm(y_k1)
#         lambda_k = x_k @ y_k1.T
#         if abs(np.linalg.norm(a @ x_k1) - np.linalg.norm(lambda_k * x_k1)) <= eps:
#             break
#         x_k = np.copy(x_k1)
        d * 
    
```

```python
# Note that the B-N-S vector $q$ looks to be a better approximation for the eigenvector.

d, u = generate_d_u(6, rndm=np.random.default_rng(1234))
A = np.diag(d) + u[:, None] @ u[None, :]
la, x = power_iter(A, x0=np.ones(6), niter=30)

q = BNS(d, u, la)

print(np.linalg.norm(la * x - A @ x))
print(np.linalg.norm(la * q - A @ q))
```

<!-- grade_id: cell-75affba958c9af75, points: 2, grade: true, solution: false -->
```python
q1 = BNS(d1, u1, la1)
assert np.linalg.norm(la1 * q1 - A1 @ q1) < 1e-9

```

## Implement the the matrix-free version of the above prototype


First, implement the function which computes the result of the matrix-vector product of the matrix 

$$\left(D + \dfrac{u\, u^T}{u^T u} \right)^{-1}$$

and a given vector $x$. Your implementation must never construct the matrix explicitly. Use the [Sherman-Morrison formula](https://en.wikipedia.org/wiki/Sherman%E2%80%93Morrison_formula).

<!-- grade_id: cell-2bb178411dfcf57c, points: 0, grade: false, solution: true -->
```python
def SMproduct(d, u, x):
    """Compute the matrix-vector product $(D + u u^T)^{-1} x$.
    
    Here $D$ is the diagonal matrix with `d` on the main diagonal,
    and `u` is a 1D array representing the (column) vector $u$.
    
    `x` is also a 1D array representing the $x$ vector.
    
    Use the Sherman-Morrison formula.
    
    Parameters
    ----------
    d : ndarray of floats, shape (n,)
        $D$ is the diagonal matrix with `d` on the main diagonal
    u : ndarray of floats, shape (n,)
        a 1D array representing the vector $u$
    x : ndarray of floats, shape (n,)
        a 1D array representing the vector $x$
    
    Returns
    -------
    y : array-like of floats, shape (n,)
        The result of the matrix-vector product of $D+(u u^T)^{-1} x$
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-59d9b1b7ecd28a9f, points: 2, grade: true, solution: false -->
```python
# Test your matrix-vector product routine on small size matrices

n = 5
rndm = np.random.default_rng(1234)

d, u = generate_d_u(n, rndm)
A = np.diag(d) + u[:, None] @ u[None, :]

for x in [np.arange(n), np.ones(n), rndm.uniform(size=n)]:
    assert_allclose(SMproduct(d, u, x),
                    np.linalg.solve(A, x), atol=1e-10)
```

## Compute the leading eigenvalue of $(D + u u^T)^{-1}$.

Write a function which computes the leading eigenvalue and its corresponding eigenvector.

<!-- grade_id: cell-fd3839d046b3752c, points: 0, grade: false, solution: true -->
```python
def leading_eigv(d, u, x0, eps=1e-7, maxiter=1000):
    """Compute the leading eigenvalue of (D + u@u.T)^{-1}.
    
    Use the power iteration method to compute the eigenvalue.
    
    Use the B-N-S formula to compute the corresponding eigenvector.
    
    Parameters
    ----------
    d : array, shape (n,)
        The diagonal elements
    u : array, shape (n,)
        The outer product vector
    x0 : array, shape (n,)
        The initial approximation
    eps : float
        The target accuracy.
        Iterations stop when |true_lambda - current estimate| < eps.
    maxiter : int
        The maximum allowed number of iterations.
    
    Returns
    -------
    lmbda : float
        The estimate for the leading eigenvalue
    vec : array, shape (n,)
        The estimate for the corresponding eigenvector
        
    Raises
    ------
    ValueError
        if the maximum number of iterations `maxiter` is reached
        without reaching the target tolerance.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-9147b698f61621a7, points: 2, grade: true, solution: false -->
```python
# test on a small example (useful for debugging)
n = 5
rndm = np.random.default_rng(1234)
d, u = generate_d_u(n, rndm)

la1, x = leading_eigv(d, u, np.ones(n))
assert np.linalg.norm(SMproduct(d, u, x) - la1*x) < 1e-9

e = eigvals(np.diag(d) + u[:, None] @ u[None, :])
e = e[np.argsort(abs(e))]
assert_allclose(e[0], 1./la1, atol=1e-14, rtol=0)
```

<!-- grade_id: cell-47d20ae241d1646a, points: 1, grade: true, solution: false -->
```python
# test interface
import pytest
with pytest.raises(ValueError):
    la1, x = leading_eigv(d, u, np.ones(n), eps=1e-20, maxiter=1)
```

<!-- grade_id: cell-21df48a03479fdf2, points: 10, grade: true, solution: false -->
```python
# test on a slightly larger example
n = 100
rndm = np.random.default_rng(1234)
d, u = generate_d_u(n, rndm)

la1, x = leading_eigv(d, u, np.ones(n))
assert np.linalg.norm(SMproduct(d, u, x) - la1*x) < 1e-9

# and a yet larger one
n = int(1e5)
rndm = np.random.default_rng(1234)
d, u = generate_d_u(n, rndm)

la1, x = leading_eigv(d, u, np.ones(n))
assert np.linalg.norm(SMproduct(d, u, x) - la1*x) < 1e-9
```