Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
import numpy as np

NAME = "Ilya Grebnekin"
COLLABORATORS = "-"
```

---

# Two-point Gauss quadrature

Construct a Gaussian quadrature formula which integrates exactly cubic polynomials over the interval $[a, b]$ with the weighting function $w(x) = 1$. Note that it is enough to use a _two-point_ quadrature formula:

$$
\int_a^b\! f(x)\, dx \approx w_1 f(x_1) + w_2 f(x_2)
$$

Write a function which returns the nodes nodes, $x_1$ and $x_2$, and weights, $w_1$ and $w_2$, of the two-point Gaussian quadrature with the weighting function $w(x) = 1$ on $[a, b]$.

```python
import numpy as np
```

<!-- grade_id: cell-9794df4167f2e72d, points: 0, grade: false, solution: true -->
```python
def gauss_2(a, b):
    """Return nodes and weights for a two-point Gauss quadrature on [a, b].

    Parameters
    ----------
    a, b : float
       Boundaries of the integration interval

    Returns
    -------
    x1, x2, w1, w2 : float
       Nodes and weights of the quadrature.
    """

    c1, c2 = (b - a) / 2, (b + a) / 2
    x1, x2 = -c1 / np.sqrt(3) + c2, c1 / np.sqrt(3) + c2

    w1, w2 = c1, c1

    return x1, x2, w1, w2
```

<!-- grade_id: cell-e9111ed2725bd5e2, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

x1, x2, w1, w2 = gauss_2(0, 1)

def f(x, n):
    return x**n

for n in [0, 1, 2, 3]:
    assert_allclose(w1*f(x1, n=n) + w2*f(x2, n=n),
                    1./(n+1), atol=1e-14)
```

```python

```

```python

```