Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ilya Grenekin"
COLLABORATORS = "-"
```

---

# Интеграл от осциллирующей функции

Рассмотрим определенный интеграл

$$
I = \int_0^\infty\! \frac{\sin(x) \cos{(\cos{(x)})}}{x}\,dx
$$

Вычислите значение данного интеграла с относительной точностью $10^{-6}$. Для упрощения задачи можете воспользоваться функционалом `scipy.integrate.quad`.

Заметим, что "из коробки" интеграл вычислить не удается, и нужно что-то придумать.

```python
from math import sin, cos

from scipy.integrate import quad

quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```

Output:
```
C:\Users\acer\AppData\Local\Temp\ipykernel_3408\270789437.py:4: IntegrationWarning: The integral is probably divergent, or slowly convergent.
  quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
(1.9653912540956733, 4.089174284042322)
```

Напишите функцию, которая возвращает значение данного интеграла. Только само значение интеграла, без оценки погрешности.
(оборачиваем интеграл в функцию только для удобства автопроверки).

```python
def f(x):
    if x == 0:
        return cos(cos(x))
    else:
        return sin(x) * cos(cos(x)) / x

def fsin(y, z):
    if y == 0:
        return 0
    else:
        return sin(z/y)*y**3

def trape(f, n,  ):
    h = 1/2 - 0
    I = (I1(0) + I2(0) + I1(1/2) + I2(1/2)) / 2
    for i in range (1, npts-1):
        I += I1(i*(h/npts))
        I += I2(i*(h/npts))
    I /= npts

    I -= (1.625 + 0.375)
```

<!-- grade_id: cell-ddd4e0b6d148fcfd, points: 0, grade: false, solution: true -->
```python
from scipy.integrate import quad
from scipy import special

def integ():
    """Return a single float, the computed value of the integral."""
    I = quad(f, 0, 1)[0]
    I += f(1)

    for k in range(0, 4):
        I -= 2 * (-1) ** k * special.jn(2 * k + 1, 1) * (
                    -sin(2 * k + 1) / (2 * k + 1) + 2 * cos(2 * k + 1) / (2 * k + 1) ** 2 + 6 * sin(2 * k + 1) / (
                        2 * k + 1) ** 3 - (24 / (2 * k + 1) ** 3) * trape(fsin, 5000, 2 * k + 1))

    return I
```

<!-- grade_id: cell-b341c2ce36bc957d, points: 5, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

# this is a test cell, keep it intact

```

<!-- grade_id: cell-b1efbd0432fcd7cb, points: 15, grade: true, solution: false -->
```python
# this is a test cell, keep it intact

```