Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ilya Grebenkin"
COLLABORATORS = "-"
```

```python
import numpy as np
```

---

# Integrable singularities

Compue the integral

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx.
$$

using the trapezoid formula. Note that the integrand diverges at the integration limits and the singularities are integrable (the integral itself exists). So, you will need to subtract these singularities.

The following integrals might be useful:

$$
\int_0^1 \frac{1}{\sqrt{x (1-x)}} \, dx=\pi,\quad \int_0^1 \frac{x}{\sqrt{x (1-x)}} \, dx=\pi/2.
$$

Transform the integral to subtract the singularities. Write out the formulas in the cell below (5pt, manual grading).

$$
\int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx = \int_{0}^{1/2}\frac{e^x}{\sqrt{x(1-x)}}dx + \int_{1/2}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx = \int_{0}^{1/2}\frac{e^x - (1+x+\frac{x^2}{2})}{\sqrt{x(1-x)}}dx + \int_{0}^{1/2}\frac{1+x+\frac{x^2}{2}}{\sqrt{x(1-x)}}dx + e \cdot  \int_{0}^{1/2}\frac{e^{-y}}{\sqrt{y(1-y)}}dy = \int_{0}^{1/2}\frac{e^x - (1+x+\frac{x^2}{2})}{\sqrt{x(1-x)}}dx + \int_{0}^{1/2}\frac{1+x+\frac{x^2}{2}}{\sqrt{x(1-x)}}dx + e \cdot  \int_{0}^{1/2}\frac{e^{-y} - (1-y-\frac{y^2}{2})}{\sqrt{y(1-y)}}dy + e \cdot \int_{0}^{1/2}\frac{1-y-\frac{y^2}{2}}{\sqrt{y(1-y)}}dy
$$

```python
def I1(x):
    return (np.exp(x) - (1 + x + x**2/2)) / (x * (1 - x))**0.5

def I2(x):
    return (np.exp(-x) - (1 - x - x**2/2)) / (x * (1 - x))**0.5
```

Write a function which computes the value of the integrals using the trapezoid rule.

<!-- grade_id: cell-9e8a974c206201a3, points: 0, grade: false, solution: true -->
```python
def integ(npts=10):
    """Compute the value of the integral above.
    
    Subtract the singularities and use the trapezoid rule. 
    
    Parameters
    ----------
    npts : int
        The number of points for the trapezoid rule
        
    Returns
    -------
    I : float
       The computed value of the integral
    """
     
    h = 1/2 - 0
    I = (I1(0) + I2(0) + I1(1/2) + I2(1/2)) / 2  
    for i in range (1, npts-1):
        I += I1(i*(h/npts))
        I += I2(i*(h/npts))
    I /= npts
    
    I -= (1.625 + 0.375)
    
    return I
```

<!-- grade_id: cell-d30dfbfd2e3409f2, points: 10, grade: true, solution: false -->
```python
# this is a test to check your computed value
from numpy.testing import assert_allclose

```

```python

```